# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np

def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
    # raise NotImplementedError()
    data = pd.read_csv('assets/NISPUF17.csv')
    df = pd.DataFrame(data)
    print(df.columns)
    print(df[['EDUC1']])
    
    education_col = 'EDUC1'
    
    # Calculate the total number of records
    total_count = df[education_col].count()
    
    # Calculate the proportions
    less_than_high_school = df[df[education_col] == 1].shape[0] / total_count
    high_school = df[df[education_col] == 2].shape[0] / total_count
    more_than_high_school = df[df[education_col] == 3].shape[0] / total_count
    college = df[df[education_col] == 4].shape[0] / total_count
    
    # Create the result dictionary
    proportions = {
        "less than high school": less_than_high_school,
        "high school": high_school,
        "more than high school but not college": more_than_high_school,
        "college": college
    }
    
    return proportions

print(proportion_of_education())
```

Output:
```
Index(['Unnamed: 0', 'SEQNUMC', 'SEQNUMHH', 'PDAT', 'PROVWT_D', 'RDDWT_D',
       'STRATUM', 'YEAR', 'AGECPOXR', 'HAD_CPOX',
       ...
       'XVRCTY2', 'XVRCTY3', 'XVRCTY4', 'XVRCTY5', 'XVRCTY6', 'XVRCTY7',
       'XVRCTY8', 'XVRCTY9', 'INS_STAT2_I', 'INS_BREAK_I'],
      dtype='object', length=454)
       EDUC1
0          4
1          3
2          3
3          4
4          1
...      ...
28460      3
28461      2
28462      3
28463      2
28464      4

[28465 rows x 1 columns]
{'less than high school': 0.10202002459160373, 'high school': 0.172352011241876, 'more than high school but not college': 0.24588090637625154, 'college': 0.47974705779026877}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # YOUR CODE HERE
    # raise NotImplementedError()
    # Read the data
    data = pd.read_csv('assets/NISPUF17.csv')
    
    print(data[['CBF_01', 'P_NUMFLU']])
    data = data[data['CBF_01'].isin([1, 2])]
    
    # Calculate the average number of influenza doses for children who were breastfed
    breastfed = data[data['CBF_01'] == 1]['P_NUMFLU'].mean()
    
    # Calculate the average number of influenza doses for children who were not breastfed
    not_breastfed = data[data['CBF_01'] == 2]['P_NUMFLU'].mean()
    
    # Return the result as a tuple
    return (breastfed, not_breastfed)

print(average_influenza_doses())
```

Output:
```
CBF_01  P_NUMFLU
0           1       NaN
1           2       NaN
2           2       NaN
3           2       3.0
4           1       0.0
...       ...       ...
28460       2       NaN
28461       1       NaN
28462       1       NaN
28463       2       NaN
28464       1       NaN

[28465 rows x 2 columns]
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # Read the data
    data = pd.read_csv('assets/NISPUF17.csv')
    
    # Filter for vaccinated children (at least one varicella dose)
    vaccinated = data[data['P_NUMVRC'] >= 1]
    
    # Separate by sex
    males = vaccinated[vaccinated['SEX'] == 1]
    females = vaccinated[vaccinated['SEX'] == 2]
    
    # Calculate the ratio for males
    male_had_cpox = males[males['HAD_CPOX'] == 1].shape[0]
    male_no_cpox = males[males['HAD_CPOX'] == 2].shape[0]
    male_ratio = male_had_cpox / male_no_cpox if male_no_cpox != 0 else 0
    
    # Calculate the ratio for females
    female_had_cpox = females[females['HAD_CPOX'] == 1].shape[0]
    female_no_cpox = females[females['HAD_CPOX'] == 2].shape[0]
    female_ratio = female_had_cpox / female_no_cpox if female_no_cpox != 0 else 0
    
    # Create the result dictionary
    ratios = {
        "male": male_ratio,
        "female": female_ratio
    }
    
    return ratios

# Call the function and print the result
print(chickenpox_by_sex())

```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    # df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
    #                "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # # here is some stub code to actually run the correlation
    # corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    #return corr

    # YOUR CODE HERE
    # raise NotImplementedError()
    # Read the data
    data = pd.read_csv('assets/NISPUF17.csv')
    
    # Filter the data: 'HAD_CPOX' should be less than 3 (1: Yes, 2: No) and drop NaN values in the selected columns
    df = data[data['HAD_CPOX'] < 3].loc[:, ['HAD_CPOX', 'P_NUMVRC']].dropna()

    # Rename the columns for clarity
    df.columns = ['had_chickenpox_column', 'num_chickenpox_vaccine_column']

    # Calculate the Pearson correlation coefficient
    corr, pval = stats.pearsonr(df["had_chickenpox_column"], df["num_chickenpox_vaccine_column"])

    # Return the correlation coefficient
    return corr
print(corr_chickenpox())
```

Output:
```
0.07044873460147556
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```